# Foundations of Computer Science (2 of 4)

Fill in the various `TODO` items, then submit a copy of this notebook 
to Craig Ferguson (`me@craigfe.io`) at least 48 hours before the supervision. 


```OCaml
let crsid = "gtf23"

;; assert (crsid <> "TODO")
```

Output:
```
val crsid : string = "gtf23"

- : unit = ()
```

<hr/>

## 1. List utility functions

Define the following utility functions on lists. Any use of recursion should be
tail-recursive. State the space complexity of each function in terms of its
non-functional parameters.

- `reverse` $[x_1; ...; x_n]$ is $[x_n; ...; x_1]$;

<!-- grade_id: cell-29fe5c84badbeb57, points: 0, grade: false, solution: true -->
```OCaml
let rec reverse_imp l acc =
    match l with
    |[] -> acc
    |x::xs -> reverse_imp xs (x::acc)
    
let reverse l = reverse_imp l []
```

Output:
```
val reverse_imp : 'a list -> 'a list -> 'a list = <fun>

val reverse : 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-3550ab5b3e9baa1c, points: 1, grade: true, solution: false -->
```OCaml
(* Ensure that [reverse] has the correct type *)
;; let __ : 'a. 'a list -> 'a list = reverse in ()

;; assert (reverse [] = [])
;; assert (reverse [ 1 ] = [ 1 ])
;; assert (reverse [ 1; 2 ] = [ 2; 1 ])
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

<!-- grade_id: cell-d2fde9c2a22aa778, points: 1, grade: true, solution: true -->
Time complexity: linear in length
Space: linear as well

- `map` $[x_1; ...; x_n]$ is $[f(x_1); ...; f(x_n)]$;

<!-- grade_id: cell-fa3c3aa002dd0562, points: 0, grade: false, solution: true -->
```OCaml
let rec map_imp f l acc =
match l with
    |[]->reverse acc
    |x::xs -> map_imp f xs ((f x)::acc)

let map f l = 
    map_imp f l []
```

Output:
```
val map_imp : ('a -> 'b) -> 'a list -> 'b list -> 'b list = <fun>

val map : ('a -> 'b) -> 'a list -> 'b list = <fun>
```

<!-- grade_id: cell-1d0e14511ae5421f, points: 1, grade: true, solution: false -->
```OCaml
;; let __ : 'a 'b. ('a -> 'b) -> 'a list -> 'b list = map in ()
;; assert (map (fun _ -> assert false) [] = [])
;; assert (map ((+) 10) [1; 2; 3] = [11; 12; 13])
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```

<!-- grade_id: cell-075a19a9925d7924, points: 1, grade: true, solution: true -->
Time: linear (times f complexity)
Space: linear

- `filter` $p\ l$ is the list containing the elements of $l$ that satisfy the predicate $p$, in the order that they appear in $l$.

<!-- grade_id: cell-8ec68ac8560c1b39, points: 0, grade: false, solution: true -->
```OCaml
let rec filter_imp f l acc =
match l with
    |[]->reverse acc
    |x::xs -> filter_imp f xs (if f x then x::acc else acc)

let filter f l =
    filter_imp f l []
```

Output:
```
val filter_imp : ('a -> bool) -> 'a list -> 'a list -> 'a list = <fun>

val filter : ('a -> bool) -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-6f40be3250ea274c, points: 1, grade: true, solution: false -->
```OCaml
;; let __ : 'a. ('a -> bool) -> 'a list -> 'a list = filter in ()
;; assert (filter (fun _ -> assert false) [] = [])
;; assert (filter (fun _ -> false) [1; 2; 3] = [])
;; assert (filter ((>) 0) [1; -1; 2; -2] = [-1; -2])
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

<!-- grade_id: cell-dce843d68b8b3478, points: 1, grade: true, solution: true -->
Linear time and space complexity. Note: we could use a special map here in filter, which assignes the value itself if the predicate is satisfied, and a constructor "None" otherwise. After this, we can remove all None-s from the list.

- `fold_left` $f\ a\ [b_1; b_2; ...; b_n]$ is the accumulated value $a'$ resulting from adding each element $b_i$ to the initial accumulator $a$ (from $b_1$ to $b_n$) with the binary operator $f$. That is: $f\ (...\ (f\ (f\ a\ b_1)\ b_2)\ ...)\ b_n$.

<!-- grade_id: cell-9c0a03457dd87de3, points: 0, grade: false, solution: true -->
```OCaml
let rec fold_left f a = function
    |[] -> a
    |x::xs -> fold_left f (f a x) xs
```

Output:
```
val fold_left : ('a -> 'b -> 'a) -> 'a -> 'b list -> 'a = <fun>
```

<!-- grade_id: cell-31e8baf606625294, points: 1, grade: true, solution: false -->
```OCaml
;; let __ : 'a 'b. ('a -> 'b -> 'a) -> 'a -> 'b list -> 'a = fold_left in ()
;; assert (fold_left (fun _ -> assert false) true [])
;; assert (fold_left (fun acc x -> x - acc) 0 [1; 2; 3] = 2)
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```

<!-- grade_id: cell-f105c45681c500cc, points: 1, grade: true, solution: true -->
Time complexity: linear.
Space complexity: also linear.

<hr/>

## 2. Partial application

For this question, you are highly encouraged to use the list utility functions
you defined above.


- `increment_all` $[x_1;\ ...;\ x_n]$ is $[x_1 + 1;\ ...;\ x_n + 1]$;

<!-- grade_id: cell-257fa9b36f742c88, points: 0, grade: false, solution: true -->
```OCaml
let increment_all = map ((+) 1)
```

Output:
```
val increment_all : int list -> int list = <fun>
```

<!-- grade_id: cell-30d32d91aba20bf1, points: 1, grade: true, solution: false -->
```OCaml
;; assert (increment_all [] = [])
;; assert (increment_all [ 1; 2; 4 ] = [ 2; 3; 5 ])
```

Output:
```
- : unit = ()

- : unit = ()
```

- `sum` $[x_1; ...; x_n]$ is $\sum_{i=1}^n x_i$;

<!-- grade_id: cell-2a75ddc9ddcc1fe1, points: 0, grade: false, solution: true -->
```OCaml
let sum = fold_left (+) 0
```

Output:
```
val sum : int list -> int = <fun>
```

<!-- grade_id: cell-07f8729a8010bd78, points: 1, grade: true, solution: false -->
```OCaml
;; assert (sum [] = 0)
;; assert (sum [1; 2; 3] = 6)
```

Output:
```
- : unit = ()

- : unit = ()
```

- `pairs_sum_to` $n\ l$ returns the list of all _unique ordered pairs_ $(x, y)$ such that $x, y \in l$ and $x + y = n$.


<!-- grade_id: cell-6ca2747158dc81de, points: 0, grade: false, solution: true -->
```OCaml
let get_unique = function  
```

<!-- grade_id: cell-27adfef709823527, points: 4, grade: true, solution: false -->
```OCaml
(* Set equality as given by structural comparison *)
let ( =~ ) a b = List.sort compare a = List.sort compare b

;; assert (pairs_sum_to 5 [1; 4]       =~ [(1, 4); (4, 1)])
;; assert (pairs_sum_to 5 [1; 2; 3; 4] =~ [(1, 4); (2, 3); (3, 2); (4, 1)])
;; assert (pairs_sum_to 5 [1; 1; 1; 4] =~ [(1, 4); (4, 1)])
;; assert (pairs_sum_to 5 [1; 1; 4; 4] =~ [(1, 4); (4, 1)])
;; assert (pairs_sum_to 2 [1; 1; 1; 1] =~ [(1, 1)])
```

<hr/>


## 3. Trees

Consider the following type of polymorphic trees:


```OCaml
type 'a tree = Branch of 'a * 'a tree list | Leaf
```

Output:
```
type 'a tree = Branch of 'a * 'a tree list | Leaf
```


1. Provide a function `prune` that takes a predicate $p$ and a tree $t$ and
   returns $t$ with any branches `Branch` $(v, ts)$ such that
   $p(v) =$ `false` replaced with `Leaf` nodes.


<!-- grade_id: cell-b7b595060a6aa4dd, points: 0, grade: true, solution: true -->
```OCaml
let rec prune p = function
    |Leaf -> Leaf
    |Branch(v, tf) -> if p v then Branch(v, map (prune p) tf) else Leaf
```

Output:
```
val prune : ('a -> bool) -> 'a tree -> 'a tree = <fun>
```

<!-- grade_id: cell-d76787056801f179, points: 4, grade: true, solution: false -->
```OCaml
;; let __ : 'a. ('a -> bool) -> 'a tree -> 'a tree = prune in ()
;; assert (prune ((=) 0) (Branch (1, [ Branch (0, []); Leaf ])) = Leaf)
;; assert (prune ((=) 0) (Branch (0, [ Branch (0, []); Branch (1, []) ])) = Branch (0, [ Branch (0, []); Leaf ]))
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```

2. Provide a function `count_subleaves` that takes a tree and replaces each value
   in the tree with a count of the number of `Leaf` nodes beneath that value in
   the tree.


<!-- grade_id: cell-ebeed3688957d1ab, points: 0, grade: false, solution: true -->
```OCaml
let rec count_subleaves = function 
    |Leaf -> Leaf
    |Branch(v, tl) -> 
        let subtree = map count_subleaves tl in
        let leafsum = sum (map (function |Branch(v, tl)->v |Leaf ->1) subtree) in
        Branch(leafsum, subtree)
```

Output:
```
val count_subleaves : 'a tree -> int tree = <fun>
```

<!-- grade_id: cell-a36eb99d6cc88726, points: 4, grade: true, solution: false -->
```OCaml
;; let __ : 'a. 'a tree -> int tree = count_subleaves in ()
;; assert (count_subleaves Leaf = Leaf)
;; assert (count_subleaves (Branch ((), [])) = Branch (0, []))
;; assert (count_subleaves 
      (Branch ((), [Branch ((), [ Leaf; Leaf ]); Leaf; Leaf ]))
    = (Branch (4,  [Branch (2,  [ Leaf; Leaf ]); Leaf; Leaf ]))
   )
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

<hr/>

## 4. Binary search tree

The polymorphic comparison operators provided by the OCaml standard library (`<`, `<=`, `=` etc.) define one specific order for each type, but this is not the only plausable ordering. For instance, we might want to sort a list of integers by their _absolute_ values:

```OCaml
let compare_abs x y = compare (abs x) (abs y)

;; List.sort compare_abs [ 1; 2; 3; 0; -1; -2; -3 ]
```

Output:
```
val compare_abs : int -> int -> int = <fun>

- : int list = [0; 1; -1; 2; -2; 3; -3]
```

For this reason, functions such as `List.sort` that require types to have some notion of _order_ often take an explicit parameter that defines that order (and using the polymorphic comparison operators is generally considered bad practice unless you _know_ that the ordering that it provides is correct). An obvious way to define orderings on a type is via a _comparator function_ on that type.

An `'a comparator` is a function $c$ that defines an ordering on values of
type `'a` such that:

- `c x y = Less_than` when $x < y$
- `c x y = Equal` when $x = y$
- `c x y = Greater_than` when $x > y$


```OCaml
type    ordering   = Less_than | Equal | Greater_than
type 'a comparator = 'a -> 'a -> ordering
```

Output:
```
type ordering = Less_than | Equal | Greater_than

type 'a comparator = 'a -> 'a -> ordering
```

One data-structure that requires a notion of ordered types is a _binary search tree_:

```OCaml
type 'a bst = Branch of 'a * 'a bst * 'a bst | Leaf
```

Output:
```
type 'a bst = Branch of 'a * 'a bst * 'a bst | Leaf
```

An `'a` binary search tree is an `'a bst` and an `'a comparator`,
$c$, such that for each branch, $\texttt{Branch}\ (v,\ l,\ r)$, in the tree:

- all values in $l$ are _less than_ $v$ as defined by $c$;
- all values in $r$ are _greater than_ $v$ as defined by $c$;
- there are no values in either $l$ or $r$ that are _equal to_ $v$ as defined by $c$.

Define the following functions on binary search trees:

1. `insert` $c\ v$ is a function that inserts the value $v$ into trees
   that are well-ordered with respect to comparator $c$. (2 marks)

<!-- grade_id: cell-244307f376007166, points: 0, grade: false, solution: true -->
```OCaml
let rec insert c v = function
    |Leaf -> Branch(v, Leaf, Leaf)
    |Branch(w, t1, t2) -> 
    if (c v w = Equal) then Branch(w, t1, t2) else
        if (c v w = Greater_than) then insert c v t1 else insert c v t2
```

Output:
```
val insert : ('a -> 'a -> ordering) -> 'a -> 'a bst -> 'a bst = <fun>
```

<!-- grade_id: cell-db34da7edc23794b, points: 2, grade: true, solution: false -->
```OCaml
;; let __ : 'a. 'a comparator -> 'a -> 'a bst -> 'a bst = insert in ()
```

Output:
```
- : unit = ()
```


2. `fold f acc t` walks the binary tree `t` and returns accumulates the
    elements into the accumulator `acc` using the function `f`, just as with
    the fold functions on lists. The order in which elements are accumulated
    does not matter.


<!-- grade_id: cell-1653b1a9d9cd075a, points: 0, grade: false, solution: true -->
```OCaml
let rec fold f acc = function
    |Leaf -> acc
    |Branch(v, t1, t2) -> 
    let right = fold f (f acc v) t1 in
    fold f right t2
```

Output:
```
val fold : ('a -> 'b -> 'a) -> 'a -> 'b bst -> 'a = <fun>
```

<!-- grade_id: cell-97e28ec79d48bb6e, points: 2, grade: true, solution: false -->
```OCaml
;; let __ : 'a 'b. ('a -> 'b -> 'a) -> 'a -> 'b bst -> 'a = fold in ()
;; assert (fold (+) 0 Leaf = 0)
;; assert (fold (+) 0 (Branch (1, Branch (2, Leaf, Leaf), Branch (3, Leaf, Leaf))) = 6)
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()
```


A `map` function on binary trees takes three parameters:

  - a function $f$ of type $\alpha \rightarrow \alpha$;

  - a comparator $c$ over values of type $\alpha$;

  - a binary tree $t$ with values of type $\alpha$, which is well-ordered with
    respect to $c$.

and returns a binary search tree containing the values $f(v)$ for each $v$
in $t$, that is also well-ordered with respect to $c$. 
Provide two implementations of `map`:

3. `map_monotonic`, which assumes that $f$ is order-preserving with
  respect to comparator $c$. That is: $\forall x, y \in \alpha.\ c\ x\ y$ = $c\ (f\ x)\ (f\ y)$.

4. `map_nonmonotonic`, which does not make this assumption.


<!-- grade_id: cell-23f4a231237ba7ac, points: 0, grade: false, solution: true -->
```OCaml
let rec map_monotonic f c = function
    |Leaf -> Leaf
    |Branch(v, t1, t2) -> Branch(f v, map_monotonic f c t1, map_monotonic f c t2)
```

Output:
```
val map_monotonic : ('a -> 'b) -> 'c -> 'a bst -> 'b bst = <fun>
```

<!-- grade_id: cell-94ff5527531bf042, points: 2, grade: true, solution: false -->
```OCaml
;; let __ : 'a. 'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst = map_monotonic in ()
```

Output:
```
File "[139]", line 1, characters 66-79:
1 | ;; let __ : 'a. 'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst = map_monotonic in ()
                                                                      ^^^^^^^^^^^^^
Error: This expression has type
         'a comparator -> ('a -> 'a) -> 'a bst -> ('a -> ordering) bst
       but an expression was expected of type
         'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst
       The type variable 'a occurs inside 'a -> ordering
```

<!-- grade_id: cell-752e1345c21448b0, points: 0, grade: false, solution: true -->
```OCaml
let rec map_nonmonotonic f c=
    let swap_arg_insert newmap v = insert c (f v) newmap in
    fold swap_arg_insert Leaf
```

Output:
```
val map_nonmonotonic :
  ('a -> 'a) -> ('a -> 'a -> ordering/2) -> 'a bst -> 'a bst = <fun>
```

<!-- grade_id: cell-4fd7f0891b2de808, points: 3, grade: true, solution: false -->
```OCaml
;; let __ : 'a. 'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst = map_nonmonotonic in ()
```

Output:
```
File "[137]", line 1, characters 66-82:
1 | ;; let __ : 'a. 'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst = map_nonmonotonic in ()
                                                                      ^^^^^^^^^^^^^^^^
Error: This expression has type
         ('a -> 'a) -> ('a -> 'a -> ordering/2) -> 'a bst -> 'a bst
       but an expression was expected of type
         'a comparator -> ('a -> 'a) -> 'a bst -> 'a bst
       Type 'a -> 'a is not compatible with type
         'a comparator = 'a -> 'a -> ordering/1 
       The type variable 'a occurs inside 'a -> ordering/1
       File "[126]", line 1, characters 0-53:
         Definition of type ordering/1
       File "[60]", line 1, characters 0-53:
         Definition of type ordering/2
```

<hr/>

## 5. Extension questions (_non-compulsory_)

`fold` functions are functions for _summarising_ data structures. `fold_right`
is identical to `fold_left`, but consumes elements of the list from _right to
left_. That is:

> `fold_right` $f\ [b_1; b_2; ... b_n]\ a$ = $f\ b_1\ (f\ b_2\ (...\ (f\ b_n\ a)\ ...))$

Note that this is equivalent to replacing every `( :: )` node in the list with $f$, and every `[]` node (of which there is only one) with $a$.


1. implement `fold_right`.

<!-- grade_id: cell-65dfef2de82985cb, points: 1, grade: true, solution: true -->
```OCaml
let fold_right : 'a 'b. ('b -> 'a -> 'a) -> 'b list -> 'a -> 'a =
  failwith "TODO: solution"
```


2. Implement all of the list utility functions in terms of a _single_ call to
   `fold_right`. Your functions should **not** be recursive or define any inner
   recursive functions.


<!-- grade_id: cell-89b346a36c0d909c, points: 1, grade: true, solution: true -->
```OCaml
let length : 'a. 'a list -> int = fun l ->
  fold_right
    failwith "TODO: solution"
```

<!-- grade_id: cell-72b665e1f11ef32b, points: 1, grade: true, solution: true -->
```OCaml
let map : 'a 'b. ('a -> 'b) -> 'a list -> 'b list = fun f l ->
  fold_right
    failwith "TODO: solution"
```

<!-- grade_id: cell-195b25a1f6e6ba28, points: 1, grade: true, solution: true -->
```OCaml
let filter : 'a. ('a -> bool) -> 'a list -> 'a list = fun p l ->
  fold_right
    failwith "TODO: solution"
```

<!-- grade_id: cell-9893537b1f1fb47c, points: 1, grade: true, solution: true -->
```OCaml
let append : 'a. 'a list -> 'a list -> 'a list = fun l1 l2 ->
  fold_right
    failwith "TODO: solution"
```

<!-- grade_id: cell-29c7df02d9c8fc5a, points: 3, grade: true, solution: true -->
```OCaml
(* N.B. this one is harder *)
let fold_left : 'a 'b. ('a -> 'b -> 'a) -> 'a -> 'b list -> 'a = fun f a l ->
  fold_right
    failwith "TODO: solution"
```