# Foundations of Computer Science (3 of 4)

Fill in the various `TODO` items, then submit a copy of this notebook 
to Craig Ferguson (`me@craigfe.io`) at least 48 hours before the supervision. 


## A. Functional queues

(_a_) Define a type `queue` that represents queues with arbitrary element type. (1 mark)

<!-- grade_id: cell-ecee44e526c730dc, points: 1, grade: true, solution: true -->
```OCaml
type 'a queue = ('a list) * ('a list)
```

Output:
```
type 'a queue = 'a list * 'a list
```

  
(_b_) Define functions `enqueue` and `dequeue` that add and remove elements from a
   queue respectively. The functions should have amortized $O(1)$ time complexity.
   State the amortized space complexity of your implementations. (2 marks each)
                                                                

<!-- grade_id: cell-f81106caf631fcd9, points: 3, grade: true, solution: true -->
```OCaml
exception Empty_queue

let norm = function
    |([], back) -> (List.rev back, [])
    |q -> q

let enqueue q v=
    match q with
    |(front, back) -> norm (front, v::back)

let dequeue = function
    |(x::front, back) -> norm (front, back)
    |_ -> raise Empty_queue
    
let qhd = function
    |(v::tail, _)->v
    
```

Output:
```
exception Empty_queue

val norm : 'a list * 'a list -> 'a list * 'a list = <fun>

val enqueue : 'a list * 'a list -> 'a -> 'a list * 'a list = <fun>

val dequeue : 'a list * 'a list -> 'a list * 'a list = <fun>

File "[11]", line 15, characters 10-39:
15 | ..........function
16 |     |(v::tail, _)->v
Warning 8: this pattern-matching is not exhaustive.
Here is an example of a case that is not matched:
([], _)
val qhd : 'a list * 'b -> 'a = <fun>
```

<!-- grade_id: cell-56f527e129503014, points: 1, grade: true, solution: true -->
Amortized space complexity: I'm not entirely sure. We input an n long list in each case to the function, but that is not on the stack. I think List.rev operates on the stack, but similarly to the time complexity, that is an amortized constant space.

## B. Sorting lists

For each of the following sorting algorithms, give an implementation for OCaml lists. (4 marks each)

These sorting functions should be parameterised by the comparator functions we used last time:

```OCaml
type    ordering   = Less_than | Equal | Greater_than
type 'a comparator = 'a -> 'a -> ordering
```

Output:
```
type ordering = Less_than | Equal | Greater_than

type 'a comparator = 'a -> 'a -> ordering
```

**Merge sort** takes an unsorted list and repeatedly applies the step of merging adjacent sorted sub-components until the entire list is sorted.

> ```ocaml
> val merge_sort : 'a comparator -> 'a list -> 'a list
> ```

<!-- grade_id: cell-f7a194a2223d4662, points: 0, grade: false, solution: true -->
```OCaml
let rec take_n l n= 
    match n, l with
    |(_,[]) -> []
    |(0,_)->[]
    |(_,x::xs) -> x::take_n xs (n-1)


let rec comb c l1 l2 =
        match (l1, l2) with
        |([],_) -> l2
        |(_,[]) -> l1
        |(l1h::l1t, l2h::l2t) -> if c l1h l2h = Less_than then l1h::(comb c l1t l2) else l2h::(comb c l1 l2t)

let rec merge_sort c l =
    match l with
    |[]->[]
    |[x]->[x]
    |x::xs->
    let front = take_n l ((List.length l)/2) in
    let back  = take_n (List.rev l) ((List.length l)/2 + ((List.length l) mod 2)) in
    let ordered_front = merge_sort c front in
    let ordered_back = merge_sort c back in
    comb c ordered_front ordered_back
```

Output:
```
val take_n : 'a list -> int -> 'a list = <fun>

val comb : ('a -> 'a -> ordering) -> 'a list -> 'a list -> 'a list = <fun>

val merge_sort : ('a -> 'a -> ordering) -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-d08b389dc07de696, points: 4, grade: true, solution: false -->
```OCaml
<`TODO`: _answer_><`TODO`: _answer_><`TODO`: _answer_>(* Input / output test cases for merge sort *)

module Test_merge = struct
  let int a b = match compare a b with
    | n when n < 0 -> Less_than
    | 0            -> Equal
    | _            -> Greater_than

  let ( => ) a b = merge_sort int a = b
end

;; let __ : 'a. 'a comparator -> 'a list -> 'a list = merge_sort in ()
;; assert Test_merge.([]                 => [])
;; assert Test_merge.([1]                => [1])
;; assert Test_merge.([2; 1]             => [1; 2])
;; assert Test_merge.([6; 1; 4; 4; 2; 3] => [1; 2; 3; 4; 4; 6])
;; assert Test_merge.(
       [ 16;  5;  6;  4; 12;  1; 15;  2;  9; 11;  6; 11; 15;  3;  9; 16;  8; 12;  6; 14; 20; 16; 18; 17;  4
       ; 13; 19; 12;  4; 11; 17; 20; 16;  9; 20;  3; 11;  3;  2; 19; 18; 13;  6; 15;  2;  4; 20; 14;  2;  4 ]
                        
    => [  1;  2;  2;  2;  2;  3;  3;  3;  4;  4;  4;  4;  4;  5;  6;  6;  6;  6;  8;  9;  9;  9; 11; 11; 11
       ; 11; 12; 12; 12; 13; 13; 14; 14; 15; 15; 15; 16; 16; 16; 16; 17; 17; 18; 18; 19; 19; 20; 20; 20; 20 ]
   )
```

Output:
```
File "[58]", line 1, characters 0-1:
1 | <`TODO`: _answer_><`TODO`: _answer_><`TODO`: _answer_>(* Input / output test cases for merge sort *)
    ^
Error: Syntax error
```

```OCaml
(* Input / output test cases for merge sort *)

module Test_merge = struct
  let int a b = match compare a b with
    | n when n < 0 -> Less_than
    | 0            -> Equal
    | _            -> Greater_than

  let ( => ) a b = merge_sort int a = b
end

;; let __ : 'a. 'a comparator -> 'a list -> 'a list = merge_sort in ()
;; assert Test_merge.([]                 => [])
;; assert Test_merge.([1]                => [1])
;; assert Test_merge.([2; 1]             => [1; 2])
;; assert Test_merge.([6; 1; 4; 4; 2; 3] => [1; 2; 3; 4; 4; 6])
;; assert Test_merge.(
       [ 16;  5;  6;  4; 12;  1; 15;  2;  9; 11;  6; 11; 15;  3;  9; 16;  8; 12;  6; 14; 20; 16; 18; 17;  4
       ; 13; 19; 12;  4; 11; 17; 20; 16;  9; 20;  3; 11;  3;  2; 19; 18; 13;  6; 15;  2;  4; 20; 14;  2;  4 ]
                        
    => [  1;  2;  2;  2;  2;  3;  3;  3;  4;  4;  4;  4;  4;  5;  6;  6;  6;  6;  8;  9;  9;  9; 11; 11; 11
       ; 11; 12; 12; 12; 13; 13; 14; 14; 15; 15; 15; 16; 16; 16; 16; 17; 17; 18; 18; 19; 19; 20; 20; 20; 20 ]
   )
```

Output:
```
module Test_merge :
  sig
    val int : 'a -> 'a -> ordering
    val ( => ) : 'a list -> 'a list -> bool
  end

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

**Quick sort** repeatedly performs 'passes' consisting of selecting a pivot element $p$, partitioning the remaining elements according to their ordering relative to $p$, then sorting each resulting sub-list as a sub-problem.


> ```ocaml
> val quick_sort : 'a comparator -> 'a list -> 'a list
> ```

<!-- grade_id: cell-2f7cbe5b929b4e70, points: 0, grade: false, solution: true -->
```OCaml
let rec quick_sort_imp c l appendTo=
    match l with
    |[]-> appendTo
    |[x]-> x::appendTo
    |x::xs->
    let smaller = List.filter (fun y -> (<>) Less_than (c x y)) xs in
    let greater = List.filter (fun y -> (=) Less_than (c x y)) xs in
    let ordered_greater = x::(quick_sort_imp c greater appendTo) in   
    quick_sort_imp c smaller ordered_greater

let rec quick_sort c l = quick_sort_imp c l []
```

Output:
```
val quick_sort_imp : ('a -> 'a -> ordering) -> 'a list -> 'a list -> 'a list =
  <fun>

val quick_sort : ('a -> 'a -> ordering) -> 'a list -> 'a list = <fun>
```

<!-- grade_id: cell-a29d38ca059435da, points: 4, grade: true, solution: false -->
```OCaml
(* Input / output test cases for quick sort *)

module Test_quick = struct
  let ( => ) a b = quick_sort Test_merge.int a = b
end

;; let __ : 'a. 'a comparator -> 'a list -> 'a list = quick_sort in ()
;; assert Test_quick.([]                 => [])
;; assert Test_quick.([1]                => [1])
;; assert Test_quick.([2; 1]             => [1; 2])
;; assert Test_quick.([6; 1; 4; 4; 2; 3] => [1; 2; 3; 4; 4; 6])
;; assert Test_quick.(
       [ 16;  5;  6;  4; 12;  1; 15;  2;  9; 11;  6; 11; 15;  3;  9; 16;  8; 12;  6; 14; 20; 16; 18; 17;  4
       ; 13; 19; 12;  4; 11; 17; 20; 16;  9; 20;  3; 11;  3;  2; 19; 18; 13;  6; 15;  2;  4; 20; 14;  2;  4 ]
        
    => [  1;  2;  2;  2;  2;  3;  3;  3;  4;  4;  4;  4;  4;  5;  6;  6;  6;  6;  8;  9;  9;  9; 11; 11; 11
       ; 11; 12; 12; 12; 13; 13; 14; 14; 15; 15; 15; 16; 16; 16; 16; 17; 17; 18; 18; 19; 19; 20; 20; 20; 20 ]
   )
```

Output:
```
module Test_quick : sig val ( => ) : 'a list -> 'a list -> bool end

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

## C. Functional arrays

This question concerns _'functional arrays'_ (Braun trees) as defined in the lecture notes. The following type represents a functional array:

```OCaml
<`TODO`: _answer_>type 'a braun_tree = Elt of 'a * 'a braun_tree * 'a braun_tree | Void
```

Output:
```
File "[17]", line 1, characters 0-1:
1 | <`TODO`: _answer_>type 'a braun_tree = Elt of 'a * 'a braun_tree * 'a braun_tree | Void
    ^
Error: Syntax error
```

```OCaml
type 'a braun_tree = Elt of 'a * 'a braun_tree * 'a braun_tree | Void
```

Output:
```
type 'a braun_tree = Elt of 'a * 'a braun_tree * 'a braun_tree | Void
```

where `Void` denotes the empty functional array $[]$ and `Elt (v, l, r)` denotes the functional array $[v; x_0; ...; x_{2n \pm 1}]$, such that:
 - `l` denotes the functional array $[x_0; x_2; x_4; ...; x_{2n}]$;
 - `r` denotes the functional array $[x_1; x_3; x_5; ...; x_{2n \pm 1}]$.

<hr/>
 
Write a function `init : int -> (int -> 'a) -> 'a braun_tree` such that:

> `init` $n\ f = [f\ 0;\ f\ 1;\ ...;\ f\ (n-1)]$,

given that $n$ is non-negative. (5 marks)


<!-- grade_id: cell-1f17bb61ea0435da, points: 0, grade: false, solution: true -->
```OCaml
let rec build i l n f= 
    if i>n then Void else
    let left = build (i+l) (2*l) n f in
    let right= build (i+2*l) (2*l) n f in
    Elt(f (i-1), left, right)
    
let init = build 1 1
(*Why does this not work? *)

let init n f = build 1 1 n f
```

Output:
```
val build : int -> int -> int -> (int -> 'a) -> 'a braun_tree = <fun>

val init : int -> (int -> '_weak1) -> '_weak1 braun_tree = <fun>

val init : int -> (int -> 'a) -> 'a braun_tree = <fun>
```

<!-- grade_id: cell-1bd19a72b5b418ad, points: 5, grade: true, solution: false -->
```OCaml
;; let __ : 'a. int -> (int -> 'a) -> 'a braun_tree = init in ()
;; assert (init 0 Fun.id = Void)
;; assert (init 1 Fun.id = Elt (0, Void, Void))
;; assert (init 2 Fun.id = Elt (0, Elt (1, Void, Void), Void))
;; assert (init 3 Fun.id = Elt (0, Elt (1, Void, Void), Elt (2, Void, Void)))
;; assert (init 4 Fun.id = Elt (0, Elt (1, Elt (3, Void, Void), Void), Elt (2, Void, Void)))
;; assert (init 5 Fun.id = Elt (0, Elt (1, Elt (3, Void, Void), Void), Elt (2, Elt (4, Void, Void), Void)))
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

Write a function `pop : 'a braun_tree -> 'a braun_tree` such that:

> `pop` $[] = []$
>
> `pop` $[x_0; ...; x_n] = [x_1; ...; x_n]$

(5 marks)

<!-- grade_id: cell-6db226171d261ded, points: 0, grade: false, solution: true -->
```OCaml
let rec pop = function
    |Void -> Void
    |Elt(v, Void, _) -> Void
    |Elt(v, left, right) ->
    match left with |Elt(w, _, _) -> Elt(w, right, pop left)
```

Output:
```
File "[21]", line 5, characters 4-60:
5 |     match left with |Elt(w, _, _) -> Elt(w, right, pop left)
        ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
Warning 8: this pattern-matching is not exhaustive.
Here is an example of a case that is not matched:
Void
val pop : 'a braun_tree -> 'a braun_tree = <fun>
```

<!-- grade_id: cell-7df5c9b67a4d5dc0, points: 5, grade: true, solution: false -->
```OCaml
;; assert (pop Void = Void)
;; assert (pop (Elt (0, Void, Void)) = Void)

;; assert (pop (Elt (0, Elt (1, Void, Void), Void))
              = Elt (1, Void, Void))

;; assert (pop (Elt (0, Elt (1, Void, Void), Elt (2, Void, Void))) 
              = Elt (1, Elt (2, Void, Void), Void))

;; assert (pop (Elt (0, Elt (1, Elt (3, Void, Void), Void), Elt (2, Void, Void))) 
              = Elt (1, Elt (2, Void, Void),                Elt (3, Void, Void)))
              
;; assert (pop (Elt (0, Elt (1, Elt (3, Void, Void), Void), Elt (2, Elt (4, Void, Void), Void))) 
              = Elt (1, Elt (2, Elt (4, Void, Void), Void), Elt (3, Void, Void)))
```

Output:
```
- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()

- : unit = ()
```

## D. Lazy data structurees

The remainder of the supervision work will concern the following three lazy data
structures. Note that each one may be used to encode a data structure of
infinite size.


```OCaml
type 'a lazy_list  = Cons   of 'a * (unit -> 'a lazy_list) | Nil
type 'a lazy_tree  = Branch of 'a * (unit -> 'a lazy_tree) * (unit -> 'a lazy_tree)
type 'a lazy_graph = Vertex of 'a * (unit -> 'a lazy_graph) list
```

Output:
```
type 'a lazy_list = Cons of 'a * (unit -> 'a lazy_list) | Nil

type 'a lazy_tree =
    Branch of 'a * (unit -> 'a lazy_tree) * (unit -> 'a lazy_tree)

type 'a lazy_graph = Vertex of 'a * (unit -> 'a lazy_graph) list
```


Write a function,

> ```ocaml
> val ndeep : int -> 'a lazy_tree -> 'a lazy_list
> ```

such that '`ndeep` $n\ t$' is a lazy list of all the $2^n$ values at
depth exactly $n$ in the tree (from left-to-right). (10 marks)


<!-- grade_id: cell-e7756b6cc180b5c6, points: 0, grade: false, solution: true -->
```OCaml
let rec traverse appendix n t =
    match (n, t) with
    |(0, Branch(v, _, _)) -> Cons(v, fun ()->appendix)
    |(n, Branch(_, left, right))->
    let rightlist = traverse appendix (n-1) (right ()) in
    traverse rightlist (n-1) (left ())

let ndeep n t = traverse Nil n t
```

Output:
```
val traverse : 'a lazy_list -> int -> 'a lazy_tree -> 'a lazy_list = <fun>

val ndeep : int -> 'a lazy_tree -> 'a lazy_list = <fun>
```

<!-- grade_id: cell-e2ca3b69f052de75, points: 10, grade: true, solution: false -->
```OCaml
module Test_tree = struct
  let rec seq n = 
    let next () = seq (n + 1) in
    Branch (n, next, next)
    
  let rec force_list = function
    | Nil -> []
    | Cons (x, xf) -> x :: force_list (xf ())
end

;; let __ : 'a. int -> 'a lazy_tree -> 'a lazy_list = ndeep in ()
;; assert Test_tree.(force_list (ndeep 3 (seq 0)) = [3; 3; 3; 3; 3; 3; 3; 3])

```

Output:
```
module Test_tree :
  sig
    val seq : int -> int lazy_tree
    val force_list : 'a lazy_list -> 'a list
  end

- : unit = ()

- : unit = ()
```

You are given an `int lazy_tree` and told that it contains arbitrarily large values at least somewhere in it. You want to find a value from it that is bigger than a given integer $n$ (but if there are many such values it does not matter which is returned). Note that because the tree is infinite, you cannot use simple depth-first search in order to visit all nodes in the tree. Present and explain code that searches the lazy tree. (8 marks)


> ```ocaml
> val find_greater_than : int -> int lazy_tree -> int
> ```

<!-- grade_id: cell-75b100ae624e2bd4, points: 0, grade: false, solution: true -->
```OCaml
let rec level levelList =
    let rec iterate this nextLevel = 
        match this with
        |[]->level nextLevel
        |Branch(v, t1, t2)::xs-> Cons(v, fun ()->(iterate xs ((t1 ())::(t2 ())::nextLevel)))
    in
    iterate levelList []
    
let find_greater_than i t = 
    let lazyList = level [t] in
    let rec search p =
        match p with
        |Cons(v, t)->if v > i then v else search (t ())
    in
    search lazyList
```

Output:
```
val level : 'a lazy_tree list -> 'a lazy_list = <fun>

File "[26]", line 12, characters 8-76:
12 | ........match p with
13 |         |Cons(v, t)->if v > i then v else search (t ())
Warning 8: this pattern-matching is not exhaustive.
Here is an example of a case that is not matched:
Nil
val find_greater_than : 'a -> 'a lazy_tree -> 'a = <fun>
```

<!-- grade_id: cell-3d94a967421f732b, points: 8, grade: true, solution: false -->
```OCaml
;; assert (find_greater_than 7 (Test_tree.seq 0) = 8)
```

Output:
```
- : unit = ()
```

Write a function `traverse` that takes a lazy graph $g$ and returns a lazy list $l$ that contains all of the elements in $g$.

> ```ocaml
> val traverse : 'a lazy_graph -> 'a lazy_list
> ```

<!-- grade_id: cell-712beb2b097a95e5, points: 0, grade: false, solution: true -->
```OCaml
failwith "TODO: solution"
```

<!-- grade_id: cell-e9f27332003e654b, points: 0, grade: true, solution: false -->
```OCaml
;; let __ : 'a. 'a lazy_graph -> 'a lazy_list = traverse in ()
```

###  Extension: lazy matrices (_non-compulsory_)

Write a function `diag` that takes an _infinite_ lazy list of _infinite_ lazy lists (a lazy matrix),

> ```
> [[ z_11; z_12; z_13; ... ];
>  [ z_21; z_22; z_23; ... ];
>  [ z_31; z_32; z_33; ... ]; ...]
> ```

and returns the diagonal: the infinite lazy list `[ z_11; z_22; z_33; ... ]`. (3 marks)

<!-- grade_id: cell-1b3d3404d35c2b7a, points: 3, grade: true, solution: true -->
```OCaml
failwith "TODO: solution"
```

Write a function `cross_init` that takes two infinite lazy lists $[ x_1; x_2; ...]$ and $[ y_1; y_2; ... ]$ and a function $f$ of two arguments; it should return a lazy matrix as above, with $z_{ij} = f\ x_i\ y_j$. (3 marks)

> ```ocaml
> val cross_init : 'a lazy_list -> 'b lazy_list -> ('a -> 'b -> 'c) -> 'c lazy_list
> ```

<!-- grade_id: cell-4242bd5eda7cfa2b, points: 3, grade: true, solution: true -->
```OCaml
failwith "TODO: solution"
```

Write a function that converts a lazy matrix as above to an infinite lazy list whose elements are all of the $z_{ij}$ in some order. (8 marks)

> ```ocaml
>   val flatten : 'a lazy_list lazy_list -> 'a lazy_list
```

<!-- grade_id: cell-52be70961132b550, points: 8, grade: true, solution: true -->
```OCaml
failwith "TODO: solution"
```